Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Mongkol  Boondamnoen"
STUDENT_ID = "65070186"
```

---

# WEEK 08: Array Oriented with NumPy - Exercise

**Write you codes after `# YOUR CODE HERE`**

<mark>**DO NOT IMPORT EXTERNAL LIBRARY**</mark>

- In **CrosswordPuzzle.pdf** is a puzzle word game
- All 10x10 puzzle characters and all 9-solutions are given below
- Use this puzzle words in the exercise
- Hints: You MAY use the following techniques (not necessary)
    - To convert from string to list of characters, try `list('abc')` and you will get `['a', 'b', 'c']`
    - To convert from a list of characters to string, try `''.join(['a', 'b', 'c'])` and you will get `'abc'`

```python
# WORDS
row0 = 'LAHEACOMET'
row1 = 'IZMYSELAOB'
row2 = 'FUFOTYRICH'
row3 = 'EGSCRDOSUN'
row4 = 'STAROSCAPT'
row5 = 'AMNENOKNLO'
row6 = 'BOAQALEHAI'
row7 = 'MOLMUKTOND'
row8 = 'PNZLTAMEES'
row9 = 'WYSPACELTY'

# SOLUTIONS WORDS
sol0 = 'SPACE'
sol1 = 'STAR'
sol2 = 'UFO'
sol3 = 'COMET'
sol4 = 'SUN'
sol5 = 'ASTRONAUT'
sol6 = 'MOON'
sol7 = 'ROCKET'
sol8 = 'PLANET'
```

---
## <mark>Execise #1 (Total: 30 points) </mark>
Create the array of puzzle-words from *CrosswordPuzzle.pdf*
- Use the **NumPy Array** to store all 10x10 characters
    - The result is an array that contains the characters in *CrosswordPuzzle.pdf*
- Name the `array` as `puzzle`

<!-- grade_id: cell-59600047be4b741a, points: 0, grade: false, solution: true -->
```python
# 1. Create NumPy Array: Crossword Puzzle 10x10
import numpy as np

data = np.array([row0, row1, row2, row3, row4, row5, row6, row7, row8, row9])
data_list2D = []

for txt in data:
    data_list = []
    for che in txt:
        data_list.append(che)
    data_list2D.append(data_list)

puzzle = np.array(data_list2D)

# YOUR CODE HERE
```

```python
# TEST. You should get these results

puzzle.ndim             # 2 <-- the puzzle is an 2D array 
puzzle.shape            # (10, 10) <-- the puzzle has 10 rows and 10 columns
print(puzzle[0])        # ['L' 'A' 'H' 'E' 'A' 'C' 'O' 'M' 'E' 'T'] <-- the row #0 (all columns)
print(puzzle[:, 3])     # ['E' 'Y' 'O' 'C' 'R' 'E' 'Q' 'M' 'L' 'P'] <-- the column #3 (all rows)
```

Output:
```
['L' 'A' 'H' 'E' 'A' 'C' 'O' 'M' 'E' 'T']
['E' 'Y' 'O' 'C' 'R' 'E' 'Q' 'M' 'L' 'P']
```

<mark>DO NOT CHANGE THE FOLLOWING LINES</mark>

Automatic Grading with the following tests. The correct results will get total **30 points**

<!-- grade_id: cell-20844b16452ccaca, points: 3, grade: true, solution: false -->
```python
# check 1: 3 points
assert(sum(puzzle[0] == list(row0)) == 10)
```

<!-- grade_id: cell-3a56ae3ae65e2086, points: 3, grade: true, solution: false -->
```python
# check 2: 3 points
assert(sum(puzzle[1] == list(row1)) == 10)
```

<!-- grade_id: cell-bf4ad5e7688fb784, points: 3, grade: true, solution: false -->
```python
# check 3: 3 points
assert(sum(puzzle[2] == list(row2)) == 10)
```

<!-- grade_id: cell-b58a04b0c3803159, points: 3, grade: true, solution: false -->
```python
# check 4: 3 points
assert(sum(puzzle[3] == list(row3)) == 10)
```

<!-- grade_id: cell-4caf5cf62adafbe1, points: 3, grade: true, solution: false -->
```python
# check 5: 3 points
assert(sum(puzzle[4] == list(row4)) == 10)
```

<!-- grade_id: cell-fdf2199ecb2c96dc, points: 3, grade: true, solution: false -->
```python
# check 6: 3 points
assert(sum(puzzle[5] == list(row5)) == 10)
```

<!-- grade_id: cell-6f5eadfca7ac58bf, points: 3, grade: true, solution: false -->
```python
# check 7: 3 points
assert(sum(puzzle[6] == list(row6)) == 10)
```

<!-- grade_id: cell-93e1145844f9c6a2, points: 3, grade: true, solution: false -->
```python
# check 8: 3 points
assert(sum(puzzle[7] == list(row7)) == 10)
```

<!-- grade_id: cell-ee1e7339435e4653, points: 3, grade: true, solution: false -->
```python
# check 9: 3 points
assert(sum(puzzle[8] == list(row8)) == 10)
```

<!-- grade_id: cell-a043b533c41cca79, points: 3, grade: true, solution: false -->
```python
# check 10: 3 points
assert(sum(puzzle[9] == list(row9)) == 10)
```

---
## <mark>Execise #2 (Total: 70 points) </mark>
Create a `find_index()` that searches the index of the given solution (`key_word`) from the `puzzle`.
- The `find_index()` returns four indexes
    - `a` is the **starting index** of a **row** that contains the `key_word`
    - `b` is the **ending index** of a **row** that contains the `key_word`
    - `c` is the **starting index** of a **column** that contains the `key_word`
    - `d` is the **ending index** of a **column** that contains the `key_word`
- For example#1, if the `key_word` is `SPACE` that is a word from the `puzzle` last row (`WYSPACELTY`)
    - The `find_index()` returns `a=9, b=10, c=2, d=7`
    - Giving the `puzzle[a:b, c:d]` means `puzzle[9:10, 2:7]` returns `SPACE`
        - The `SPACE` is located at `puzzle`'s row `9` and column `2:7`
- For example#2, if the `key_word` is `ASTRONAUT` that is a word from the `puzzle` column `4`
    - The `find_index()` returns `a=0, b=9, c=4, d=5`
        - Giving the `puzzle[a:b, c:d]` means `puzzle[0:9, 4:5]` returns `ASTRONAUT`
            - The `ASTRONAUT` is located at `puzzle`'s column `4` and row `0:9`

<!-- grade_id: cell-9d3217f89a689f42, points: 0, grade: false, solution: true -->
```python
# 2. Search index of a key word
def find_index(puzzle, key_word):
    # YOUR CODE HERE
    row_puzzle = np.array([row0, row1, row2, row3, row4, row5, row6, row7, row8, row9])
    che_puzzle = puzzle.T
    column_puzzle = []
    che = False
    for txt in che_puzzle:
        text = ""
        for txt_2 in txt:
            text += txt_2
        column_puzzle.append(text)
    for mes_row in range(len(row_puzzle)):
        if row_puzzle[mes_row].find(key_word) != -1:
            a = mes_row
            b = mes_row + 1
            c = int(row_puzzle[mes_row].find(key_word))
            d = c + len(key_word)
            che = True
            break
    if che == False:
        for mes_column in range(len(column_puzzle)):
            if column_puzzle[mes_column].find(key_word) != -1:
                a = int(column_puzzle[mes_column].find(key_word))
                b = a + len(key_word)
                c = mes_column
                d = mes_column + 1
                che = True
                break
    return a, b, c, d
```

<!-- grade_id: cell-15746ee0af791ab4, points: 0, grade: true, solution: false -->
```python
# TEST 1
a, b, c, d = find_index(puzzle, sol0)       # a=9, b=10, c=2, d=7
assert(sol0 == ''.join(puzzle[a:b, c:d].flatten()))
print(f'Keyword {sol0} found in puzzle\'s row[{a}:{b}] and column[{c}:{d}]]')
```

Output:
```
Keyword SPACE found in puzzle's row[9:10] and column[2:7]]
```

```python
# TEST 2
a, b, c, d = find_index(puzzle, sol5)       # a=0, b=9, c=4, d=5
assert(sol5 == ''.join(puzzle[a:b, c:d].flatten()))
print(f'Keyword {sol5} found in puzzle\'s row[{a}:{b}] and column[{c}:{d}]]')
```

Output:
```
Keyword ASTRONAUT found in puzzle's row[0:9] and column[4:5]]
```

<mark>DO NOT CHANGE THE FOLLOWING LINES</mark>

Automatic Grading with the following tests. The correct results will get total **70 points**

<!-- grade_id: cell-5ca5be398f15f844, points: 10, grade: true, solution: false -->
```python
# check-1: 10 points
a, b, c, d = find_index(puzzle, sol1)
assert(sol1 == ''.join(puzzle[a:b, c:d].flatten()))
print(f'Keyword {sol1} found in puzzle\'s row[{a}:{b}] and column[{c}:{d}]]')
```

Output:
```
Keyword STAR found in puzzle's row[4:5] and column[0:4]]
```

<!-- grade_id: cell-6ea6d6df81d29199, points: 10, grade: true, solution: false -->
```python
# check-2: 10 points
a, b, c, d = find_index(puzzle, sol2)
assert(sol2 == ''.join(puzzle[a:b, c:d].flatten()))
print(f'Keyword {sol2} found in puzzle\'s row[{a}:{b}] and column[{c}:{d}]]')
```

Output:
```
Keyword UFO found in puzzle's row[2:3] and column[1:4]]
```

<!-- grade_id: cell-f2ee007bee1d1018, points: 10, grade: true, solution: false -->
```python
# check-3: 10 points
a, b, c, d = find_index(puzzle, sol3)
assert(sol3 == ''.join(puzzle[a:b, c:d].flatten()))
print(f'Keyword {sol3} found in puzzle\'s row[{a}:{b}] and column[{c}:{d}]]')
```

Output:
```
Keyword COMET found in puzzle's row[0:1] and column[5:10]]
```

<!-- grade_id: cell-a5042b313ff24aaa, points: 10, grade: true, solution: false -->
```python
# check-4: 10 points
a, b, c, d = find_index(puzzle, sol4)
assert(sol4 == ''.join(puzzle[a:b, c:d].flatten()))
print(f'Keyword {sol4} found in puzzle\'s row[{a}:{b}] and column[{c}:{d}]]')
```

Output:
```
Keyword SUN found in puzzle's row[3:4] and column[7:10]]
```

<!-- grade_id: cell-06b35ad376959e6d, points: 10, grade: true, solution: false -->
```python
# check-5: 10 points
a, b, c, d = find_index(puzzle, sol6)
assert(sol6 == ''.join(puzzle[a:b, c:d].flatten()))
print(f'Keyword {sol6} found in puzzle\'s row[{a}:{b}] and column[{c}:{d}]]')
```

Output:
```
Keyword MOON found in puzzle's row[5:9] and column[1:2]]
```

<!-- grade_id: cell-e18b4425b23a7b5d, points: 10, grade: true, solution: false -->
```python
# check-6: 10 points
a, b, c, d = find_index(puzzle, sol7)
assert(sol7 == ''.join(puzzle[a:b, c:d].flatten()))
print(f'Keyword {sol7} found in puzzle\'s row[{a}:{b}] and column[{c}:{d}]]')
```

Output:
```
Keyword ROCKET found in puzzle's row[2:8] and column[6:7]]
```

<!-- grade_id: cell-fa7da355c6333d5a, points: 10, grade: true, solution: false -->
```python
# check-7: 10 points
a, b, c, d = find_index(puzzle, sol8)
assert(sol8 == ''.join(puzzle[a:b, c:d].flatten()))
print(f'Keyword {sol8} found in puzzle\'s row[{a}:{b}] and column[{c}:{d}]]')
```

Output:
```
Keyword PLANET found in puzzle's row[4:10] and column[8:9]]
```