Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Mongkol  Boondamnoen"
STUDENT_ID = "65070186"
```

---

# String - File and Exception

**Write you codes after `# YOUR CODE HERE`**

Use the **provided files**

<mark>**DO NOT IMPORT EXTERNAL LIBRARY**</mark>

---
## <mark>Execise #1 **Count Words in text file** (Total: 40 points) </mark>
- Create a `read_and_count()` function that reads a text file to:
    - Count the word from the text file
    - Return the number of word
- For example, 
    - `read_and_count('song.txt', 'green')` return `1` means `green` occurred only once

<!-- grade_id: cell-4db7be5c9fcfb010, points: 0, grade: false, solution: true -->
```python
def read_and_count(file, word):
    # YOUR CODE HERE
    num = 0
    with open(file, 'r') as myfile:
        list_txt = myfile.read().split()
        for txt in list_txt:
            if txt.lower() == word:
                num += 1
        return num
```

```python
# TEST
print(read_and_count('song.txt', 'green'))  # 1
print(read_and_count('song.txt', 'you'))    # 3
```

Output:
```
1
3
```

<mark>DO NOT CHANGE THE FOLLOWING LINES</mark>

Automatic Grading with the following tests. The correct results will get total **40 points**

<!-- grade_id: cell-3b69e9135c90a845, points: 10, grade: true, solution: false -->
```python
# check-1: 10 points
assert(read_and_count('song.txt', 'world') == 4)
```

<!-- grade_id: cell-a93c1843b6e09c32, points: 10, grade: true, solution: false -->
```python
# check-2: 10 points
assert(read_and_count('song.txt', 'of') == 5)
```

<!-- grade_id: cell-5867bf2141f75299, points: 10, grade: true, solution: false -->
```python
# check-3: 10 points
assert(read_and_count('song.txt', 'the') == 6)
```

<!-- grade_id: cell-15a44fd437c40429, points: 10, grade: true, solution: false -->
```python
# check-4: 10 points
assert(read_and_count('song.txt', 'i') == 11)
```

---
## <mark> Execise #2 **Count movie by genre and rating from csv** (Total: 60 points) </mark>
- Creat a `read_csv()` function to read a `csv` file that contains all moveis
    - Each movie record contains **title**, **genre**, **rating**, and others
    - Function receives `file`, `genre`, and `rating`
        - `file` is a `csv` file
        - `genre` is a movie genre
        - `rating` is a move rating
    - Function return `3` things:
        1) A number of all movie in a given movie `genre` that have movie's rating `>= rating`
        2) A highest-rate movie title of all movies from 1)
        3) A lowest-rate movie title of all movies from 1)
    - Example, `read_csv('n_movies.csv', 'Fantasy', 7.5)` return
        - `(185, 'Stranger Things', 'All of Us Are Dead')`
        - `185` is a number of all movies that are in `Fantasy` genre and have movie's rating `>= rating`
        - The **highest** rating movie from `185` movies is `Stranger Things`
        - The **lowest** rating movie from `185` movies is `All of Us Are Dead`

<!-- grade_id: cell-792c2518da990957, points: 0, grade: false, solution: true -->
```python
# n_movies.csv is a dataset from Kaggle: https://www.kaggle.com/datasets/narayan63/netflix-popular-movies-dataset

import csv
def read_csv(file, genre, rating):    
    # YOUR CODE HERE
    num = 0
    rat = []
    story = []
    with open(file, mode='r', newline='') as mycsv:
        reader = csv.reader(mycsv)
        for record in reader:
            movie = record[4].split(", ")
            try:
                movie.index(genre.title())
                float(record[5])
            except:
                pass
            else:
                if float(record[5]) >= rating:
                    rat.append(float(record[5]))
                    story.append(record[0])
                    num += 1
    list_2d = [rat, story]
    hightrat = rat.copy()
    lowrat = rat.copy()
    hightrat = max(hightrat)
    lowrat = min(lowrat)
    hightnum = list_2d[0].index(hightrat)
    maxtxt = list_2d[1][hightnum]
    lownum = list_2d[0].index(lowrat)
    mintxt = list_2d[1][lownum]
    return num, maxtxt, mintxt
```

```python
# TEST
print(read_csv('n_movies.csv', 'Fantasy', 7.5))
# (185, 'Stranger Things', 'All of Us Are Dead')
print(read_csv('n_movies.csv', 'Crime', 8.0))
# (216, 'Dexter', 'The Blacklist')
```

Output:
```
(185, 'Stranger Things', 'All of Us Are Dead')
(216, 'Dexter', 'The Blacklist')
```

<mark>DO NOT CHANGE THE FOLLOWING LINES</mark>

Automatic Grading with the following tests. The correct results will get total **60 points**

<!-- grade_id: cell-df1b00a0172a0276, points: 20, grade: true, solution: false -->
```python
# check-1: 20 points
assert(read_csv('n_movies.csv', 'COMEDY', 8.2) == (293, 'BoJack Horseman', 'Gilmore Girls'))
```

<!-- grade_id: cell-93fd3122fdcc06f0, points: 20, grade: true, solution: false -->
```python
# check-2: 20 points
assert(read_csv('n_movies.csv', 'Thriller', 8.5) == (24, 'Breaking Bad', 'Ozark'))
```

<!-- grade_id: cell-9754868adbb97389, points: 20, grade: true, solution: false -->
```python
# check-3: 20 points
assert(read_csv('n_movies.csv', 'biography', 8.8) == (4, 'The Last Dance', 'Narcos'))
```